```python
# Initialize Otter
import otter
grader = otter.Notebook("BasicExercises.ipynb")
```

Instructions:

- Open the file in the editor or Google Colab.
- Solve each exercise by adapting the provided code skeleton.
- Each exercise is written as a function to be able to test it independently.
- Run the test cell to verify the correctness and get feedback.
- The code is tested using the `otter-grader` module.
- The test cell is hidden by default. 
- Run the last cell to submit the exercise.
- Import any required modules at the beginning of the script.

## Exercise 1: Data Types

Print data type of all variables in the provided code. First write also your guess for the data types of the variables.

The output should be like this:

```python
x = 5, My guess: int, Data type: <class 'int'>
y = 5.0, My guess: float, Data type: <class 'float'>
...
```

```python
def print_datatypes():
    a = 1
    b = 8.3
    c = "s"
    d = "word"
    e = True
    f = bool(0)
    g = None
    h = [1, 2, 3]
    i = (1, 2, 3)
    j = {1, 2, 3}
    k = {"a": 1, "b": 2, "c": 3}
    l = range(10)
    k = [1,"s",True]
    
    # BEGIN SOLUTION
    print("a = {}, My guess: int, Data type: {}".format(a, type(a)))
    print("b = {}, My guess: float, Data type: {}".format(b, type(b)))
    print("c = {}, My guess: str, Data type: {}".format(c, type(c)))
    print("d = {}, My guess: str, Data type: {}".format(d, type(d)))
    print("e = {}, My guess: bool, Data type: {}".format(e, type(e)))
    print("f = {}, My guess: bool, Data type: {}".format(f, type(f)))
    print("g = {}, My guess: None, Data type: {}".format(g, type(g)))
    print("h = {}, My guess: list, Data type: {}".format(h, type(h)))
    print("i = {}, My guess: tuple, Data type: {}".format(i, type(i)))
    print("j = {}, My guess: set, Data type: {}".format(j, type(j)))
    print("k = {}, My guess: dict, Data type: {}".format(k, type(k)))
    print("l = {}, My guess: range, Data type: {}".format(l, type(l)))
    print("m = {}, My guess: list, Data type: {}".format(m, type(m)))
    # END SOLUTION


    


```

## Exercise 2: Data Type Conversion and String Formatting

Convert the variables to the following data types:

- `a` to a float
- `b` to an integer
- `c` to a string

Think about the possible data loss during the conversion.

Print the variables after the conversion in the format `a = x.xx, b =   x, c = "x"`.



```python
def datatype_casting():
    a = 1
    b = 8.3
    c = bool(0)

    # BEGIN SOLUTION
    a = float(a)
    b = int(b)
    c = int(c)
    print("a = {%.2f}, b = \t{%d}, c = \"{%s}\"".format(a, b, c))
    # END SOLUTION

    
```

## Exercise 3: System Information - Part 1
Print Path of the Working Directory

```python
def print_working_directory():
    # BEGIN SOLUTION
    print(os.getcwd())
    # END SOLUTION
```

## Exercise 4: System Information - Part 2
List Files in the Working Directory

```python
def list_files():
    # BEGIN SOLUTION
    for file in os.listdir():
        print(file)
    # END SOLUTION
```

## Exercise 5: Counting
Count the number of characters in a string and print the result in the format `The string "xxx" has x characters.`

```python
def count_characters(s: str) -> int:
    # BEGIN SOLUTION
    print("The string '{}' has {} characters.".format(s, len(s)))
    # END SOLUTION
```

## Exercise 6: Lists
Define a list of numbers from 1 to 10, use while, print the list and return the list.

```python
def my_list():
    # BEGIN SOLUTION
    l = []
    i = 0
    while i < 10:
        l.append(i)
        i += 1
    print(l)
    return l
    # END SOLUTION
```

## Exercise 7: List Operations
1. Add 11 to the list from the previous exercise and print the list.
3. Further delete the 4th element from the list and print the list.
4. Print the length of the list.

```python
def list_operations():
    # BEGIN SOLUTION
    l = my_list()
    l.append(11)
    print(l)
    l.pop(4)
    print(l)
    print(len(l))
    # END SOLUTION

```

## Exercise 8: Loops and Condtions
Print all the numbers from 1 to 10 that are divisible by 2.


```python
def loops_conditions():
    # BEGIN SOLUTION
    for i in range(10):
        if i % 2 == 0:
            print(i)
    # END SOLUTION
```

```python
grader.check("e8")
```

## Exercise 9: List Comprehension
Create a list of squares of the numbers from 1 to 10 using list comprehension, print and return the list.


```python
def list_comprehension():
    # BEGIN SOLUTION
    l = [i for i in range(10) if i % 2 == 0]
    print(l)
    return l
    # END SOLUTION
```

## Exercise 10: Functions

Define a function `check_even` that takes a list of integer numbers as input and returns `True` if all the number in the list are even, otherwise returns `False`.

First check if the numbers are integers, if not return `None`.





```python
# BEGIN SOLUTION
def check_even():
    for i in range(10):
        if i % 2 == 0:
            print(i)
# END SOLUTION
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for e1 (points: None)
```
>>> def test_print_datatypes(capsys):
...     print_datatypes()
...     captured = capsys.readouterr()
...     output = captured.out.splitlines()
...     assert output[0] == "a = 1, My guess: int, Data type: <class 'int'>"
...     assert output[1] == "b = 8.3, My guess: float, Data type: <class 'float'>"
...     assert output[2] == "c = s, My guess: str, Data type: <class 'str'>"
...     assert output[3] == "d = word, My guess: str, Data type: <class 'str'>"
...     assert output[4] == "e = True, My guess: bool, Data type: <class 'bool'>"
...     assert output[5] == "f = False, My guess: bool, Data type: <class 'bool'>"
...     assert output[6] == "g = None, My guess: None, Data type: <class 'NoneType'>"
...     assert output[7] == "h = [1, 2, 3], My guess: list, Data type: <class 'list'>"
...     assert output[8] == "i = (1, 2, 3), My guess: tuple, Data type: <class 'tuple'>"
...     assert output[9] == "j = {1, 2, 3}, My guess: set, Data type: <class 'set'>"
...     assert output[10] == "k = {'a': 1, 'b': 2}, My guess: dict, Data type: <class 'dict'>"
...     assert output[11] == "l = range(0,10) My guess: range, Data type: <class 'range'>"
...     assert output[12] == "m = [1, 's', True], My guess: list, Data type: <class 'list'>"
```
```
>>> test_print_datatypes(capsys)
```


Tests for e10 (points: None)
```
>>> def test_check_even():
...     assert check_even([2, 4, 6]) == 'True'
...     assert check_even([1, 3, 5]) == 'False'
...     assert check_even([2, 4, 5]) == 'False'
```
```
>>> test_check_even()
```


Tests for e2 (points: None)
```
>>> def test_datatype_casting(capsys):
...     datatype_casting()
...     captured = capsys.readouterr()
...     output = captured.out.splitlines()
...     assert output[0] == 'a = {1.00}, b = \t{8}, c = "{0}"'
```
```
>>> test_datatype_casting(capsys)
```


Tests for e3 (points: None)
```
>>> def test_print_working_directory(capsys):
...     print_working_directory()
...     captured = capsys.readouterr()
...     output = captured.out.strip()
...     assert output == os.getcwd()
```
```
>>> test_print_working_directory(capsys)
```


Tests for e4 (points: None)
```
>>> def test_list_files(capsys):
...     list_files()
...     captured = capsys.readouterr()
...     output = captured.out.splitlines()
...     assert set(output) == set(os.listdir())
```
```
>>> test_list_files(capsys)
```


Tests for e5 (points: None)
```
>>> def test_count_characters():
...     assert count_characters('hello') == "The string 'hello' has 5 characters."
...     assert count_characters('') == "The string '' has 0 characters."
```
```
>>> test_count_characters()
```


Tests for e6 (points: None)
```
>>> def test_my_list():
...     my_list()
...     captured = capsys.readouterr()
...     output = captured.out.splitlines()
...     assert output == ['[0, 1, 2, 3, 4, 5, 6, 7, 8, 9]']
```
```
>>> test_my_list()
```


Tests for e7 (points: None)
```
>>> def test_list_operations(capsys):
...     list_operations()
...     captured = capsys.readouterr()
...     output = captured.out.splitlines()
...     assert output[0] == '[0, 1, 2, 3, 4, 5, 6, 7, 8, 9]'
...     assert output[1] == '[0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 11]'
...     assert output[2] == '[0, 1, 2, 3, 5, 6, 7, 8, 9, 11]'
...     assert output[3] == '10'
```
```
>>> test_list_operations()
```


Tests for e8 (points: None)
```
>>> def test_loops_conditions(capsys):
...     loops_conditions()
...     captured = capsys.readouterr()
...     output = captured.out.splitlines()
...     assert output == ['2', '4', '6', '8', '10']
```
```
>>> test_loops_conditions(capsys)
```


Tests for e9 (points: None)
```
>>> def test_list_comprehension(capsys):
...     list_comprehension()
...     captured = capsys.readouterr()
...     output = captured.out.splitlines()
...     assert output == ['[0, 2, 4, 6, 8]']
```
```
>>> test_list_comprehension()
```
