---
title: "Basics Python"
execute:
    echo: true
    eval: false
---

```python
# Initialize Otter
import otter
grader = otter.Notebook("BasicExercises.ipynb")
```

Instructions:

- Open the file in the editor or Google Colab.
- Solve each exercise by adapting the provided code skeleton.
- Each exercise is written as a function to be able to test it independently.
- Run the test cell to verify the correctness and get feedback.
- The code is tested using the `otter-grader` module.
- The test cell is hidden by default. 
- Run the last cell to submit the exercise.
- Import any required modules at the beginning of the script.

## Exercise 1: Data Types

Print data type of all variables in the provided code. First write also your guess for the data types of the variables.

The output should be like this:

```python
x = 5, My guess: int, Data type: <class 'int'>
y = 5.0, My guess: float, Data type: <class 'float'>
...
```

```python
def print_datatypes():
    a = 1
    b = 8.3
    c = "s"
    d = "word"
    e = True
    f = bool(0)
    g = None
    h = [1, 2, 3]
    i = (1, 2, 3)
    j = {1, 2, 3}
    k = {"a": 1, "b": 2, "c": 3}
    l = range(10)
    k = [1,"s",True]
    
    ...


    


```

## Exercise 2: Data Type Conversion and String Formatting

Convert the variables to the following data types:

- `a` to a float
- `b` to an integer
- `c` to a string

Think about the possible data loss during the conversion.

Print the variables after the conversion in the format `a = x.xx, b =   x, c = "x"`.



```python
def datatype_casting():
    a = 1
    b = 8.3
    c = bool(0)

    ...

    
```

## Exercise 3: System Information - Part 1
Print Path of the Working Directory

```python
def print_working_directory():
    ...
```

## Exercise 4: System Information - Part 2
List Files in the Working Directory

```python
def list_files():
    ...
```

## Exercise 5: Counting
Count the number of characters in a string and print the result in the format `The string "xxx" has x characters.`

```python
def count_characters(s: str) -> int:
    ...
```

## Exercise 6: Lists
Define a list of numbers from 1 to 10, use while, print the list and return the list.

```python
def my_list():
    ...
```

## Exercise 7: List Operations
1. Add 11 to the list from the previous exercise and print the list.
3. Further delete the 4th element from the list and print the list.
4. Print the length of the list.

```python
def list_operations():
    ...

```

## Exercise 8: Loops and Condtions
Print all the numbers from 1 to 10 that are divisible by 2.


```python
def loops_conditions():
    ...
```

```python
grader.check("e8")
```

## Exercise 9: List Comprehension
Create a list of squares of the numbers from 1 to 10 using list comprehension, print and return the list.


```python
def list_comprehension():
    ...
```

## Exercise 10: Functions

Define a function `check_even` that takes a list of integer numbers as input and returns `True` if all the number in the list are even, otherwise returns `False`.

First check if the numbers are integers, if not return `None`.





```python
...
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for e1 (points: None)


Tests for e10 (points: None)


Tests for e2 (points: None)


Tests for e3 (points: None)


Tests for e4 (points: None)


Tests for e5 (points: None)


Tests for e6 (points: None)


Tests for e7 (points: None)


Tests for e8 (points: None)
```
>>> def test_loops_conditions(capsys):
...     loops_conditions()
...     captured = capsys.readouterr()
...     output = captured.out.splitlines()
...     assert output == ['2', '4', '6', '8', '10']
```
```
>>> test_loops_conditions(capsys)
```


Tests for e9 (points: None)
